```python
# Initialize Otter
import otter
grader = otter.Notebook("lab01.ipynb")
```

```python
import matplotlib.pyplot as plt
import numpy as np
from typing import Union
%matplotlib inline
rng_seed = 42
```

# COM424 - Machine Learning
## Lab Exercise 01
Knowledge practice lab exercise on Machine Learning Foundations

**Question 1.** Define a function called `even_ints` that returns a 1D `numpy` array of even positive integers given a specified range limit. Returned array must include the limit element if it is an even positive integer.

```python
def even_ints(limit : int) -> Union[np.ndarray,None]:
    """ Generate a Numpy a array of even positive integers """
    ...

```

```python
grader.check("question1")
```

**Question 2.** Evaluate the following linear relations computationally:
$$ 7x + 6y + 3z = 20 \\
-4z + 4y + x = 14 \\
4x + y = 10$$

The solution should be a vector $u$

```python
X = ...
d = ...
```

```python
...
```

```python
...
```

```python
grader.check("q2")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q2 (points: 4)
```
>>> assert np.allclose(u, np.array([2.60176991, -0.54867257, 0.5840708]))
```


Tests for question1 (points: 3)
```
>>> from math import ceil
>>> 
>>> def test_odd_even_ints(limit):
...     assert isinstance(even_ints(limit), np.ndarray), 'Function must return a Numpy array'
...     if isinstance(even_ints(limit), np.ndarray):
...         num_elements = limit // 2 + 1
...         assert even_ints(limit).shape == tuple([limit // 2 + 1]), 'Function must return a 1D array'
```
